---
title       : Testing subexercises
description : This thing is meant to test subexercises

--- type:NormalExercise lang:r xp:100 skills:1 key:e37b77d822
## Test


*** =instructions

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sample_code
```{r}

```

*** =solution
```{r}

```

*** =sct
```{r}
success_msg("Schrödinger conqueror")
```
